---
title       : Data Types
description : We're going to get an overview, of the different types of variables that there are in R, and how to work with them. 

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:0f78ac0820
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sct
```{r}

```
